## Quick note about Jupyter cells

When you are editing a cell in Jupyter notebook, you need to re-run the cell by pressing **`<Shift> + <Enter>`** or **`<ctrl> + <Enter>`**. This will allow changes you made to be available to other cells.

Use **`<Enter>`** to make new lines inside a cell you are editing.

#### Code cells

Re-running will execute any statements you have written. To edit an existing code cell, click on it.

#### Markdown cells

Re-running will render the markdown text. To edit an existing markdown cell, double-click on it.


## Common Jupyter operations

Near the top of the Jupyter notebook window, there are a row of menu options (`File`, `Edit`, `View`, `Insert`, ...) and a row of tool bar icons (disk, plus sign, scissors, 2 files, clipboard and file, up arrow, ...).

#### Inserting and removing cells

- Use the "plus sign" icon to insert a cell below the currently selected cell
- Use "Insert" -> "Insert Cell Above" from the menu to insert above

#### Clear the output of all cells

- Use "Kernel" -> "Restart" from the menu to restart the kernel
    - click on "clear all outputs & restart" to have all the output cleared

#### Save your notebook file locally

- Clear the output of all cells
- Use "File" -> "Download as" -> "Notebook (.ipynb)"

## References

- https://jupyter-notebook.readthedocs.io/en/latest/notebook.html
- https://mybinder.readthedocs.io/en/latest/introduction.html
- https://docs.python.org/3/tutorial/index.html
- https://docs.python.org/3/tutorial/introduction.html
- https://www.rstudio.com/wp-content/uploads/2015/03/rmarkdown-reference.pdf

## Python objects, basic types, and variables

Like in R, everything in Python is an **object** and every object in Python has a **type**. Some of the basic types include:

- **`int`** (integer; a whole number with no decimal place)
  - `10`
  - `-3`
- **`float`** (float; a number that has a decimal place)
  - `7.41`
  - `-0.006`
- **`str`** (string; a sequence of characters enclosed in single quotes, double quotes, or triple quotes)
  - `'this is a string using single quotes'`
  - `"this is a string using double quotes"`
  - `'''this is a triple quoted string using single quotes'''`
  - `"""this is a triple quoted string using double quotes"""`
- **`bool`** (boolean; a binary value that is either true or false)
  - `True`
  - `False`
- **`NoneType`** (a special type representing the absence of a value)
  - `None`

In Python, a **variable** is a name you specify in your code that maps to a particular **object**, object **instance**, or value.

By defining variables, we can refer to things by names that make sense to us. Names for variables can only contain letters, underscores (`_`), or numbers (no spaces, dashes, or other characters). Variable names must start with a letter or underscore.

## Basic operators

In Python, there are different types of **operators** (special symbols) that operate on different values. Some of the basic operators include:

- arithmetic operators
  - **`+`** (addition)
  - **`-`** (subtraction)
  - **`*`** (multiplication)
  - **`/`** (division)
  - __`**`__ (exponent)
- assignment operators
  - **`=`** (assign a value)
  - **`+=`** (add and re-assign; increment)
  - **`-=`** (subtract and re-assign; decrement)
  - **`*=`** (multiply and re-assign)
- comparison operators (return either `True` or `False`)
  - **`==`** (equal to)
  - **`!=`** (not equal to)
  - **`<`** (less than)
  - **`<=`** (less than or equal to)
  - **`>`** (greater than)
  - **`>=`** (greater than or equal to)

When multiple operators are used in a single expression, **operator precedence** determines which parts of the expression are evaluated in which order. Operators with higher precedence are evaluated first (like PEMDAS in math). Operators with the same precedence are evaluated from left to right.

- `()` parentheses, for grouping
- `**` exponent
- `*`, `/` multiplication and division
- `+`, `-` addition and subtraction
- `==`, `!=`, `<`, `<=`, `>`, `>=` comparisons

> See https://docs.python.org/3/reference/expressions.html#operator-precedence

```python
# Assigning some numbers to different variables
num1 = 10
num2 = -3
num3 = 7.41
num4 = -.6
num5 = 7
num6 = 3
num7 = 11.11
```

```python
# Addition
num1 + num2
```

```python
# Subtraction
num2 - num3
```

```python
# Multiplication
num3 * num4
```

```python
# Division
num4 / num5
```

```python
# Exponent
num5 ** num6
```

```python
# Increment existing variable
num7 += 4
num7
```

```python
# Decrement existing variable
num6 -= 2
num6
```

```python
# Multiply & re-assign
num3 *= 5
num3
```

```python
# Assign the value of an expression to a variable
num8 = num1 + num2 * num3
num8
```

```python
# Are these two expressions equal to each other?
num1 + num2 == num5
```

```python
# Are these two expressions not equal to each other?
num3 != num4
```

```python
# Is the first expression less than the second expression?
num5 < num6
```

```python
# Is this expression True?
5 > 3 > 1
```

```python
# Is this expression True?
5 > 3 < 4 == 3 + 1
```

```python
# Assign some strings to different variables
simple_string1 = 'an example'
simple_string2 = "oranges "
```

```python
# Addition
simple_string1 + ' of using the + operator'
```

```python
# Notice that the string was not modified
simple_string1
```

```python
# Multiplication
simple_string2 * 4
```

```python
# This string wasn't modified either
simple_string2
```

```python
# Are these two expressions equal to each other?
simple_string1 == simple_string2
```

```python
# Are these two expressions equal to each other?
simple_string1 == 'an example'
```

```python
# Add and re-assign
simple_string1 += ' that re-assigned the original string'
simple_string1
```

```python
# Multiply and re-assign
simple_string2 *= 3
simple_string2
```

Note: **Subtraction, division, and decrement operators do not apply to strings.**

## Basic containers

> Note: **mutable** objects can be modified after creation and **immutable** objects cannot.

Containers are objects that can be used to group other objects together. The basic container types include:

- **`str`** (string: immutable; indexed by integers; items are stored in the order they were added)
- **`list`** (list: mutable; indexed by integers; items are stored in the order they were added)
  - `[3, 5, 6, 3, 'dog', 'cat', False]`
- **`tuple`** (tuple: immutable; indexed by integers; items are stored in the order they were added)
  - `(3, 5, 6, 3, 'dog', 'cat', False)`
- **`set`** (set: mutable; not indexed at all; items are NOT stored in the order they were added; can only contain immutable objects; does NOT contain duplicate objects)
  - `{3, 5, 6, 3, 'dog', 'cat', False}`
- **`dict`** (dictionary: mutable; key-value pairs are indexed by immutable keys; items are NOT stored in the order they were added)
  - `{'name': 'Jane', 'age': 23, 'fav_foods': ['pizza', 'fruit', 'fish']}`

When defining lists, tuples, or sets, use commas (,) to separate the individual items. When defining dicts, use a colon (:) to separate keys from values and commas (,) to separate the key-value pairs.

Strings, lists, and tuples are all **sequence types** that can use the `+`, `*`, `+=`, and `*=` operators.

```python
# Assign some containers to different variables
list1 = [3, 5, 6, 3, 'dog', 'cat', False]
tuple1 = (3, 5, 6, 3, 'dog', False, 'cat')
set1 = {3, 5, 6, 3, 'dog', 'cat', False}
dict1 = {'name': 'Jane', 'age': 23, 'fav_foods': ['pizza', 'fruit', 'fish']}
```

```python
# Items in the list object are stored in the order they were added
list1
```

```python
# Items in the tuple object are stored in the order they were added
tuple1
```

```python
# Items in the set object are not stored in the order they were added
# Also, notice that the value 3 only appears once in this set object
set1
```

```python
# Items in the dict object are not stored in the order they were added
dict1
```

```python
# Add and re-assign
list1 += [5, 'grapes']
list1
```

```python
# Add and re-assign
tuple1 += (5, 'grapes')
tuple1
```

```python
# Multiply
[1, 2, 3, 4] * 2
```

```python
# Multiply
(1, 2, 3, 4) * 3
```

## Accessing data in containers

For strings, lists, tuples, and dicts, we can use **subscript notation** (square brackets) to access data at an index.

- strings, lists, and tuples are indexed by integers, **starting at 0** for first item
  - these sequence types also support accesing a range of items, known as **slicing**
  - use **negative indexing** to start at the back of the sequence
- dicts are indexed by their keys

> Note: sets are not indexed, so we cannot use subscript notation to access data elements.

### Quick note about doing tasks below

replace `raise NotImplementedError()` with your code

<!-- grade_id: cell-1da7d625eb2bf9c9, points: 0, grade: false, solution: true -->
```python
# Access the first item in a sequence 'list'
#answer = list1[]
# YOUR CODE HERE
raise NotImplementedError()

answer
```

<!-- grade_id: cell-afe3407a02081c14, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 3
```

<!-- grade_id: cell-61bb9bf82f781902, points: 0, grade: false, solution: true -->
```python
#del(answer)
# Access the last item in a sequence tuple1
#answer = tuple1[]

# YOUR CODE HERE
raise NotImplementedError()
answer
```

<!-- grade_id: cell-4114ed1504558d22, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert type(answer) == str
```

<!-- grade_id: cell-62f108cba0bee8a7, points: 0, grade: false, solution: true -->
```python
del(answer)
# Access a range of items in a sequence simple_string1 to get 'examp' word
#answer = simple_string1[]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6671f272189e46b0, points: 0, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 'examp'
```

<!-- grade_id: cell-ddfbb8118abc06ea, points: 0, grade: false, solution: true -->
```python
del(answer)
# Access a range of items in a sequence from first to 3rd last
#answer = tuple1[]
# YOUR CODE HERE
raise NotImplementedError()
```

```python
"Check that you have correct result"
assert answer == (3, 5, 6, 3)
```

<!-- grade_id: cell-e17a4fe1a9415a4c, points: 0, grade: false, solution: true -->
```python
# Access last three elements from list
#answer = list1[]

# YOUR CODE HERE
raise NotImplementedError()
answer
```

<!-- grade_id: cell-7c0b32ba337d9c6a, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer[0:3] == ['dog', 'cat', False]
```

<!-- grade_id: cell-e1af0c92dd4eaae9, points: 0, grade: false, solution: true -->
```python
# Access an item in a dictionary

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3a21f83e9be84194, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 'Jane'
```

<!-- grade_id: cell-b806f05e4a772af8, points: 0, grade: false, solution: true -->
```python
# Access a third element of a sequence in a dictionary
dict1['fav_foods'][2]

# YOUR CODE HERE
raise NotImplementedError()
answer
```

```python
"Check that you have correct result"
assert answer == 'fish'
```

## Python built-in functions and callables

A **function** is a Python object that you can "call" to **perform an action** or compute and **return another object**. You call a function by placing parentheses to the right of the function name. Some functions allow you to pass **arguments** inside the parentheses (separating multiple arguments with a comma). Internal to the function, these arguments are treated like variables.

Python has several useful built-in functions to help you work with different objects and/or your environment. Here is a small sample of them:

- **`type(obj)`** to determine the type of an object
- **`len(container)`** to determine how many items are in a container
- **`callable(obj)`** to determine if an object is callable
- **`sorted(container)`** to return a new list from a container, with the items sorted
- **`sum(container)`** to compute the sum of a container of numbers
- **`min(container)`** to determine the smallest item in a container
- **`max(container)`** to determine the largest item in a container
- **`abs(number)`** to determine the absolute value of a number
- **`repr(obj)`** to return a string representation of an object

> Complete list of built-in functions: https://docs.python.org/3/library/functions.html

There are also different ways of defining your own functions and callable objects that we will explore later.

<!-- grade_id: cell-6480ccb731ca9080, points: 0, grade: false, solution: true -->
```python
# Use the type() function to determine the type of the simple_string1 object

# YOUR CODE HERE
raise NotImplementedError()
answer

```

<!-- grade_id: cell-6ef10b11e4c13c55, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == str
```

<!-- grade_id: cell-d11629e1e1baff5f, points: 0, grade: false, solution: true -->
```python
# Use the len() function to determine how many items are in a container dict1

# YOUR CODE HERE
raise NotImplementedError()
answer 

```

<!-- grade_id: cell-6d8ea0cc91e9bafa, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 3
```

<!-- grade_id: cell-07a8851718830c0f, points: 0, grade: false, solution: true -->
```python
# Use the len() function to determine how many items are in a container simple_string2

# YOUR CODE HERE
raise NotImplementedError()
answer

```

<!-- grade_id: cell-09f8f1c596a6328e, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer >7 & answer <9
```

<!-- grade_id: cell-e88eec99382f0133, points: 0, grade: false, solution: true -->
```python
# Use the callable() function to determine if an object 'len' is callable

# YOUR CODE HERE
raise NotImplementedError()

answer 
```

<!-- grade_id: cell-3e92380dad814710, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == True
```

<!-- grade_id: cell-79850c6bddfd8697, points: 0, grade: false, solution: true -->
```python
# Use the callable() function to determine if an object 'dict1' is callable

# YOUR CODE HERE
raise NotImplementedError()

answer 
```

<!-- grade_id: cell-b931643c685ddf53, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == False
```

<!-- grade_id: cell-a08188b9fd4332c4, points: 0, grade: false, solution: true -->
```python
# Use the sorted() function to return a new list from a container, with the items sorted
container = [10, 1, 3.6, 7, 5, 2, -3]
# YOUR CODE HERE
raise NotImplementedError()
answer 

```

<!-- grade_id: cell-4b96d783edbdd8e3, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == [-3, 1, 2, 3.6, 5, 7, 10]
```

```python
# Use the sorted() function to return a new list from a container, with the items sorted
# - notice that capitalized strings come first
container = ['dogs', 'cats', 'zebras', 'Chicago', 'California', 'ants', 'mice']


```

<!-- grade_id: cell-ea38175dbe666a3c, points: 0, grade: false, solution: true -->
```python
# Use the sum() function to compute the sum of a container of numbers
container = [10, 1, 3.6, 7, 5, 2, -3]

# YOUR CODE HERE
raise NotImplementedError()
answer 
```

<!-- grade_id: cell-ccf1b0454bd3cfe3, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 25.6
```

<!-- grade_id: cell-8128784db4ca3505, points: 0, grade: false, solution: true -->
```python
# Use the min() function to determine the smallest item in a container
container = [10, 1, 3.6, 7, 5, 2, -3]

# YOUR CODE HERE
raise NotImplementedError()
answer 
```

<!-- grade_id: cell-d7c726ab234a1830, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == -3
```

<!-- grade_id: cell-833e773cb74c9db5, points: 0, grade: false, solution: true -->
```python
# Use the min() function to determine the smallest item in a container
container = ['g', 'z', 'a', 'y']

# YOUR CODE HERE
raise NotImplementedError()
answer 
```

<!-- grade_id: cell-753bdf4424ee9295, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 'a'
```

<!-- grade_id: cell-8f9cc08799ad0374, points: 0, grade: false, solution: true -->
```python
# Use the max() function to determine the largest item in a container
container = [10, 1, 3.6, 7, 5, 2, -3]

# YOUR CODE HERE
raise NotImplementedError()
answer 
```

<!-- grade_id: cell-47c4b70a4a8832d6, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 10
```

<!-- grade_id: cell-bfc73068524e38d0, points: 0, grade: false, solution: true -->
```python
# Use the max() function to determine the largest item in a container
container = max('gibberish')

# YOUR CODE HERE
raise NotImplementedError()
answer 
```

<!-- grade_id: cell-3885bba1d4940fbe, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 's'
```

<!-- grade_id: cell-3281c6dd738d1d8e, points: 0, grade: false, solution: true -->
```python
# Use the abs() function to determine the absolute value of a number 5

# YOUR CODE HERE
raise NotImplementedError()
answer 
```

<!-- grade_id: cell-4373aa32982005a2, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 5
```

<!-- grade_id: cell-0c4cef9efa216ccb, points: 0, grade: false, solution: true -->
```python
# Use the abs() function to determine the absolute value of a number -12

# YOUR CODE HERE
raise NotImplementedError()
answer 
```

<!-- grade_id: cell-ca12b41ace4eaf65, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer > 0
```

## Python object attributes (methods and properties)

Different types of objects in Python have different **attributes** that can be referred to by name (similar to a variable). To access an attribute of an object, use a dot (`.`) after the object, then specify the attribute (i.e. `obj.attribute`)

When an attribute of an object is a callable, that attribute is called a **method**. It is the same as a function, only this function is bound to a particular object.

When an attribute of an object is not a callable, that attribute is called a **property**. It is just a piece of data about the object, that is itself another object.

The built-in `dir()` function can be used to return a list of an object's attributes.

## Some methods on string objects

- **`.capitalize()`** to return a capitalized version of the string (only first char uppercase)
- **`.upper()`** to return an uppercase version of the string (all chars uppercase)
- **`.lower()`** to return an lowercase version of the string (all chars lowercase)
- **`.count(substring)`** to return the number of occurences of the substring in the string
- **`.startswith(substring)`** to determine if the string starts with the substring
- **`.endswith(substring)`** to determine if the string ends with the substring
- **`.replace(old, new)`** to return a copy of the string with occurences of the "old" replaced by "new"

```python
# Assign a string to a variable
a_string = 'tHis is a sTriNg'
```

<!-- grade_id: cell-a09465f63f9fe2e1, points: 0, grade: false, solution: true -->
```python
# Return a capitalized version of the a_string

# YOUR CODE HERE
raise NotImplementedError()
answer 
```

<!-- grade_id: cell-147c5269c2d44ce0, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 'This is a string'
```

<!-- grade_id: cell-7c725e6873a0b6c0, points: 0, grade: false, solution: true -->
```python
# Return an uppercase version of the string

# YOUR CODE HERE
raise NotImplementedError()
answer 
```

<!-- grade_id: cell-485cfd4bb3e6aa35, points: 0, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 'THIS IS A STRING'
```

<!-- grade_id: cell-5285e26ac5aadb0d, points: 0, grade: false, solution: true -->
```python
# Return a lowercase version of the string

# YOUR CODE HERE
raise NotImplementedError()
answer 
```

<!-- grade_id: cell-b09f53f8b21b1d77, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 'this is a string'
```

```python
# Notice that the methods called have not actually modified the string
a_string
```

<!-- grade_id: cell-e25986f6892c47bd, points: 0, grade: false, solution: true -->
```python
# Count number of occurences of a substring 'i' in the string

# YOUR CODE HERE
raise NotImplementedError()
answer 
```

<!-- grade_id: cell-f2081a9f893a3c6c, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 3
```

```python
# Count number of occurences of a substring in the string after a certain position
a_string.count('i', 7)
```

<!-- grade_id: cell-1760fc08556c0a5d, points: 0, grade: false, solution: true -->
```python
# Count number of occurences of a substring 'is' in the string

# YOUR CODE HERE
raise NotImplementedError()
answer 
```

<!-- grade_id: cell-aab90d96ba53bde9, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 2
```

<!-- grade_id: cell-9376041d77e4c626, points: 0, grade: false, solution: true -->
```python
# Does the string start with 'this'?

# YOUR CODE HERE
raise NotImplementedError()
answer
```

<!-- grade_id: cell-1efeb5a93bf5a201, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == False
```

<!-- grade_id: cell-d4fce967ccec0fe7, points: 0, grade: false, solution: true -->
```python
# Does the lowercase string start with 'this'?

# YOUR CODE HERE
raise NotImplementedError()
answer
```

<!-- grade_id: cell-de703181f7867142, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == True
```

<!-- grade_id: cell-02f858ab105f8698, points: 0, grade: false, solution: true -->
```python
# Does the string end with 'Ng'?

# YOUR CODE HERE
raise NotImplementedError()
answer
```

<!-- grade_id: cell-29b2ae7d55f37a5b, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == True
```

<!-- grade_id: cell-2bd5d5793aea41b7, points: 0, grade: false, solution: true -->
```python
# Return a version of the string with a substring replaced with something else like ('is' for 'XYZ')

# YOUR CODE HERE
raise NotImplementedError()
answer
```

<!-- grade_id: cell-fbe440ce92f858e1, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 'tHXYZ XYZ a sTriNg'
```

<!-- grade_id: cell-b2c4143200df43bd, points: 0, grade: false, solution: true -->
```python
# Return a version of the string with a substring replaced with something else like ('i' for '!')

# YOUR CODE HERE
raise NotImplementedError()
answer
```

<!-- grade_id: cell-cc4df90a8f76417b, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 'tH!s !s a sTr!Ng'
```

<!-- grade_id: cell-8be63f3a8c26d42a, points: 0, grade: false, solution: true -->
```python
# Return a version of the string with the first 2 occurences of 'i' replaced with '!'

# YOUR CODE HERE
raise NotImplementedError()
answer
```

<!-- grade_id: cell-0eaed028fb1dae54, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 'tH!s !s a sTriNg'
```

## Some methods on list objects

- **`.append(item)`** to add a single item to the list
- **`.extend([item1, item2, ...])`** to add multiple items to the list
- **`.insert(i, item)`** to insert an item at a given position
- **`.remove(item)`** to remove a single item from the list
- **`.pop()`** to remove and return the item at the end of the list
- **`.pop(index)`** to remove and return an item at an index
- **`.count(item)`** return the number of times item appears in the list
- **`.reverse()`** to reverse order of items in the list

```python
fruits = ['orange', 'apple', 'pear', 'banana', 'kiwi', 'apple', 'banana']
```

<!-- grade_id: cell-2f877bc251ed0def, points: 0, grade: false, solution: true -->
```python
# Count apples in fruits

# YOUR CODE HERE
raise NotImplementedError()
answer
```

<!-- grade_id: cell-9ece4337dca1bbc1, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == 2
```

<!-- grade_id: cell-cecc70fdaafba009, points: 0, grade: false, solution: true -->
```python
# What's index of banana?

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-21107716c9db7957, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
```

<!-- grade_id: cell-3025e8ef6500aa1b, points: 0, grade: false, solution: true -->
```python
# Find next banana starting at position 4

# YOUR CODE HERE
raise NotImplementedError()
answer

```

<!-- grade_id: cell-837ae97db88d18c7, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
```

<!-- grade_id: cell-0801e3d7afe5230e, points: 0, grade: false, solution: true -->
```python
# Reverse fruits list

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-7faff3b3e8620340, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert fruits == ['banana', 'apple', 'kiwi', 'banana', 'pear', 'apple', 'orange']
```

<!-- grade_id: cell-a9d98f57d762cbe0, points: 0, grade: false, solution: true -->
```python
# Add grape to the list

# YOUR CODE HERE
raise NotImplementedError()
fruits
```

<!-- grade_id: cell-06f479edb807c6aa, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert fruits == ['banana', 'apple', 'kiwi', 'banana', 'pear', 'apple', 'orange', 'grape']
```

```python
fruits.sort()
fruits
```

### List comprehensions

List comprehensions provide a concise way to create lists. Common applications are to make new lists where each element is the result of some operations applied to each member of another sequence or iterable, or to create a subsequence of those elements that satisfy a certain condition.

For example, assume we want to create a list of squares, like:

```python
squares = []
for x in range(10):
    squares.append(x**2)
    
squares
```

```python
#or, equivalently:
squares = [x**2 for x in range(10)]
```

A list comprehension consists of brackets containing an expression followed by a for clause, then zero or more for or if clauses. The result will be a new list resulting from evaluating the expression in the context of the for and if clauses which follow it. For example, this listcomp combines the elements of two lists if they are not equal:

```python
[(x, y) for x in [1,2,3] for y in [3,1,4] if x != y]
```

```python
vec = [-4, -2, 0, 2, 4]
```

<!-- grade_id: cell-777c8e9e8c0245cc, points: 0, grade: false, solution: true -->
```python
# from 'vec' create a new list with the values doubled

# YOUR CODE HERE
raise NotImplementedError()
answer
```

<!-- grade_id: cell-575a1111fd262a23, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == [-8, -4, 0, 4, 8]
```

<!-- grade_id: cell-f46faa20c00cf38f, points: 0, grade: false, solution: true -->
```python
# filter the 'vec' list to exclude negative numbers

# YOUR CODE HERE
raise NotImplementedError()
answer
```

<!-- grade_id: cell-68471b9737837aa2, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == [0, 2, 4]
```

<!-- grade_id: cell-9252d605574f274e, points: 0, grade: false, solution: true -->
```python
# apply a abs function to all the elements

# YOUR CODE HERE
raise NotImplementedError()
answer
```

<!-- grade_id: cell-a80c34146b431c78, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert answer == [4, 2, 0, 2, 4]
```

That's all for today!