## Differences between `list` and numpy `ndarray'

```python
import numpy as np
```

<!-- grade_id: cell-9f97501e0d2bb2e1, points: 1, grade: true, solution: true -->
```python
# Create list named lst consisting letters from "a" to "d"

# YOUR CODE HERE
raise NotImplementedError()
print(lst)

# Create ndarray named arr same as this list
# YOUR CODE HERE
raise NotImplementedError()
print(arr)

```

<!-- grade_id: cell-69ec0617c0969596, points: 1, grade: true, solution: false -->
```python
"Check that you have correct result"
assert lst == ["a", "b", "c", "d"]
assert lst == np.ndarray.tolist(arr)
```

<!-- grade_id: length, points: 0, grade: false, solution: true -->
```python
# What is length of both?

# YOUR CODE HERE
raise NotImplementedError()

print(answer1)
print(answer2)
```

<!-- grade_id: length_test, points: 1, grade: true, solution: false -->
```python
# Hidden results check
```

<!-- grade_id: types, points: 0, grade: false, solution: true -->
```python
# What are types of objects in both?

# YOUR CODE HERE
raise NotImplementedError()

print(answer1)
print(answer2)
```

<!-- grade_id: types_test, points: 1, grade: true, solution: false -->
```python
# Hidden results check
```

<!-- grade_id: create, points: 0, grade: false, solution: true -->
```python
# Create 20 element list and ndarray consisting consecutive numbers starting from 1

# YOUR CODE HERE
raise NotImplementedError()

print( "List: ", nbr)
print("ndarray: ", arr)
```

<!-- grade_id: create_test, points: 1, grade: true, solution: false -->
```python
assert isinstance(nbr, list)
assert isinstance(arr, np.ndarray)
assert nbr[0] == 1
assert arr[0] == 1
assert nbr[-1] == 20
assert arr[-1] == 20
assert len(nbr) == 20
assert len(arr) == 20
```

<!-- grade_id: cell-7ca22235d0fa22c9, points: 0, grade: false, solution: true -->
```python
# Multiply both by 3

# YOUR CODE HERE
raise NotImplementedError()

print(answer1)
print(answer2)
```

Write below what is the difference in the results.

<!-- grade_id: roznica_mnozenia, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: add_10, points: 0, grade: false, solution: true -->
```python
# Add 10 to both

# YOUR CODE HERE
raise NotImplementedError()

print(answer1)
print(answer2)
```

Write below what is the difference in the results.

<!-- grade_id: roznica_dodawania, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: dzielenie, points: 0, grade: false, solution: true -->
```python
# Divide both by 100

# YOUR CODE HERE
raise NotImplementedError()

print(answer1) # nbr
print(answer2) # arr
```

<!-- grade_id: dzielenie_test, points: 1, grade: true, solution: false -->
```python
assert answer1 == np.ndarray.tolist(answer2)
```

## Standard Deviation Calculation

The formula for standard deviation (SD) is

$SD = \sqrt{\frac{\sum{\mid x - \mu \mid ^2}}{N}}$

where $\sum$ sum means "sum of", $x$ is a value in the data set, $\mu$ is the mean of the data set, and $N$ is the number of data points in the population.

The steps you're about to follow to calculate the standard deviation:

Step 1: Find the mean.

Step 2: For each data point, find the square of its distance to the mean.

Step 3: Sum the values from Step 2.

Step 4: Divide by the number of data points.

Step 5: Take the square root.

Our data:

```python
np.random.seed(1)
data = np.random.randn(300,2)
```

<!-- grade_id: 5_first_rows, points: 0, grade: false, solution: true -->
```python
#What are first 5 rows of the data?

# YOUR CODE HERE
raise NotImplementedError()

answer
```

<!-- grade_id: 5_first_rows_test, points: 1, grade: true, solution: false -->
```python
assert answer.ndim == 2
assert answer.size == 10

```

<!-- grade_id: sum_columns, points: 0, grade: false, solution: true -->
```python
# create new array called data1 consisting sum of these two columns

# YOUR CODE HERE
raise NotImplementedError()
data1[:10]
```

<!-- grade_id: sum_columns_test, points: 1, grade: true, solution: false -->
```python
# Hidden results check
```

<!-- grade_id: mean, points: 0, grade: false, solution: true -->
```python
# Step 1: Find the mean of data1

# YOUR CODE HERE
raise NotImplementedError()
mn
```

<!-- grade_id: mean_test, points: 1, grade: true, solution: false -->
```python
assert mn == .13092478126002602
```

<!-- grade_id: distance_mean, points: 0, grade: false, solution: true -->
```python
# Step 2: For each data point, find the square (np.square) of its distance to the mean.

# YOUR CODE HERE
raise NotImplementedError()

dist[:10]
```

<!-- grade_id: distance_mean_test, points: 1, grade: true, solution: false -->
```python
assert dist[0] == 0.777331706463166
```

<!-- grade_id: dist_sum, points: 0, grade: false, solution: true -->
```python
# Step 3: Sum the values from Step 2

# YOUR CODE HERE
raise NotImplementedError()
answer
```

<!-- grade_id: dist_sum_test, points: 1, grade: true, solution: false -->
```python
assert answer == 538.7307843276668
```

<!-- grade_id: divide_by_number, points: 0, grade: false, solution: true -->
```python
# Step 4: Divide by the number of data points

# YOUR CODE HERE
raise NotImplementedError()
answer
```

<!-- grade_id: divide_by_number_test, points: 1, grade: true, solution: false -->
```python
assert answer == 1.7957692810922228
```

<!-- grade_id: sd, points: 0, grade: false, solution: true -->
```python
# Step 5: Take the square root.
import math

# YOUR CODE HERE
raise NotImplementedError()

sd
```

<!-- grade_id: sd_test, points: 1, grade: true, solution: false -->
```python
assert sd == 1.3400631630980022
```